```python
# Initialize Otter
import otter
grader = otter.Notebook("project.ipynb")
```

# Project 01 – The Other Side of Gradescope

## DSC 80, Fall 2021

### Checkpoint Due Date: Thursday, October 7
### Due Date: Thursday, October 14

---
# Instructions

This Jupyter Notebook contains the statements of the problems and provides code and markdown cells to display your answers to the problems.  
* Like the lab, your coding work will be developed in the accompanying `project.py` file, that will be imported into the current notebook. This code will be autograded.
* **For the checkpoint, you only need to turn in a `project.py` containing solutions for questions 1-4**
    - The checkpoint autograder on Gradescope does not thoroughly check your code -- it only runs the doctests on problems 1-4 to make sure that you have completed them. When you submit the final version of the project, we will use more tests to check these answers more thoroughly.

**Do not change the function names in the `*.py` file**
- The functions in the `*.py` file are how your assignment is graded, and they are graded by their name.
- If you changed something you weren't supposed to, just use git to revert!

**Tips for developing in the .py file**:
- Do not change the function names in the starter code; grading is done using these function names.
- Do not change the docstrings in the functions. These are there to tell you if your work is on the right track!
- You are **encouraged to write your own additional functions** to solve the questions! 
    - Developing in python usually consists of larger files, with many short functions.
    - You may write your other functions in an additional `.py` file that you import in `project.py` -- however, be sure to upload these to gradescope as well!
- Always document your code!

**Tips for testing the correctness of your answers!**
Once you have your work saved in the .py file, you should import the `project` to test your function out in the notebook. In the notebook you should inspect/analyze the output to assess its correctness!
* Run your functions on the main dataset (`grades`) and ask yourself if the output *looks correct.*
* Run your functions on very small datasets (e.g. 1-5 row table), calculate the expected response by hand, and see if the function output matches (this *is* unit-testing your code with data).
* Run your functions on (large and small) samples of the dataset `grades` (with and without replacement). Does your code break? Or does it still run as expected.

```python
%load_ext autoreload
%autoreload 2
```

```python
import pandas as pd
import numpy as np
import os
```

```python
from project import *
```

## About the Assignment

The file contains the grade-book from a fictional data science course with 535 students. 

**Note: this dataset is synthetically generated; it does not contain real student grades. The course syllabus below is also similar, but not exactly the same as the course syllabus for this class!**

In this project, you will:
1. clean and process the data to compute total course grades according to a fictional syllabus (below),
2. qualitatively understand how students did in the course,
3. understand how student grades vary with small changes in performance on each assignment.

---

The course syllabus for this fictional class is as follows:

* Lab assignments 
    - Each are worth the same amount, regardless of each lab's raw point total.
    - The lowest lab is dropped.
    - Each lab may be revised for one week after submission for a 10% penalty, for two weeks after submission for a 30% penalty, and beyond that for a 60% penalty. Such revisions are reflected in the `Lateness` columns in the gradebook.
    - Labs are 20% of the total grade.
* Projects 
    - Each project consists of an autograded portion, and *possibly* a free response portion.
    - The total points for a single project consist of the sum of the raw score of the two portions.
    - Each are worth the same amount, regardless of each project's raw point total.
    - Projects are 30% of the total grade.
* Checkpoints
    - Project checkpoints are worth 2.5% of the total grade.
* Discussion
    - Discussion notebooks are worth 2.5% of the total grade.
* Exams
    - The midterm is worth 15% of the total grade.
    - The final is worth 30% of the total grade.


### A note on generalization

You may assume that your code will only need to work on a gradebook for a class with the syllabus given above. That is, you may assume that the dataframe `grades` looks like the given one in `data/grades.csv`.

However, such a class:
1. may have a different numbers of labs, projects, discussions, and project checkpoints.
2. may have a different number of students.

You may assume the course components and the naming conventions are as given in the data file.

The dataset was generated by Gradescope; you must attempt to reason about the data as given using what you know as a student who uses Gradescope.

### A note on 'putting everything together'

The goal of this project is to create and assess final grades for a fictional course; if anything, the process is broken down into functions for your convenience and guidance. Here are a few remarks and tips for approaching the projects:
1. If you are having trouble figuring out what a question is asking you to do, look at the big picture and try to understand what the current step is doing to contribute to this big picture. This may clarify what's being asked!
1. These questions intentionally build off of each other and the final result matters! In fact, you can 'get a question correct', but only receive partial credit on it because a previous answer was wrong.
    - Credit for a question will typically receive partial credit based on *how close* your answer is to correct (as well as some credit for a solution in the correct form). 
    - You should try to assess your answer to each question based on what you understand of the data. This might involve writing extensive code (that isn't turned in) just to check your work! Suggestions on checking your work are given in the assignment, but you should also think of your own ways of checking your work.
    - As you do this project, think about the data from the perspective of the student (which should be easy to do!)

```python
grades_fp = os.path.join('data', 'grades.csv')
grades = pd.read_csv(grades_fp)
```

```python
grades
```

Output:
```
PID College Level       lab01  lab01 - Max Points  \
0    A14721419      SI    JR   99.735279               100.0   
1    A14883274      TH    JR   98.829476               100.0   
2    A14164800      SI    SR   86.513369               100.0   
3    A14847419      TH    JR  100.000000               100.0   
4    A14162943      SI    JR   66.506974               100.0   
..         ...     ...   ...         ...                 ...   
530  A14490387      SI    JR  100.000000               100.0   
531  A14088257      SI    SO  100.000000               100.0   
532  A14847419      WA    JR   88.656641               100.0   
533  A14513929      TH    SR   83.799719               100.0   
534  A14365863      WA    SR   88.389331               100.0   

    lab01 - Lateness (H:M:S)       lab02  lab02 - Max Points  \
0                   00:00:00   84.990171               100.0   
1                   00:00:00   50.784231               100.0   
2                   00:00:00   47.802820          
```

### Getting started: enumerating the assignments

First, you will list all the 'assignment names' and what part of the syllabus to which they belong.

**Question 1:**

Create a function `get_assignment_names` that takes in a dataframe like `grades` and returns a dictionary with the following structure:
- The keys are the general areas of the syllabus: `lab, project, midterm, final, disc, checkpoint`
- The values are lists that contain the assignment names of that type. For example the lab assignments all have names of the form `labXX` where `XX` is a zero-padded two digit number. See the doctests for more details.

```python
def get_assignment_names(grades):
    
    conv = {'lab': 'lab', 'project': 'project',
            'Midterm': 'midterm', 'Final': 'final',
            'discussion': 'disc', 'project_checkpoint': 'checkpoint'}
    cols = ["lab", "project", "midterm", "final", "disc", "checkpoint"]
    sol = {i:list() for i in cols}

    for i in grades.columns.values:
        temp = "".join([j for j in i if not j.isdigit()])
        if temp in conv.keys():
            sol[conv[temp]].append(i)

    return sol

get_assignment_names(grades)
```

Output:
```
{'lab': ['lab01',
  'lab02',
  'lab03',
  'lab04',
  'lab05',
  'lab06',
  'lab07',
  'lab08',
  'lab09'],
 'project': ['project01', 'project02', 'project03', 'project04', 'project05'],
 'midterm': ['Midterm'],
 'final': ['Final'],
 'disc': ['discussion01',
  'discussion02',
  'discussion03',
  'discussion04',
  'discussion05',
  'discussion06',
  'discussion07',
  'discussion08',
  'discussion09',
  'discussion10'],
 'checkpoint': ['project02_checkpoint01',
  'project02_checkpoint02',
  'project03_checkpoint01']}
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

### Computing project grades

**Question 2**

Compute the total score for the project portion of the course according to the syllabus. Create a function `projects_total` that takes in `grades` and computes the total project grade for the quarter according to the syllabus. The output Series should contain values between 0 and 1.

*Note*: Don't forget to properly handle students who didn't turn in assignments! (Use your experience and common sense).

*Note:* To check your work, try (1) calculating the score for a few types of students by hand, and (2) calculate the statistics for the class performance on each individual course project, making sure they look reasonable.

```python
def projects_total(grades):
    temp = grades.fillna(0)    #replace null/None/nan values
    sol = pd.DataFrame()
    sol["PID"] = grades["PID"]
    
    for i in get_assignment_names(temp)["project"]:
        if (i + "_free_response") in temp.columns.values:
            sol[i] = ((temp[i] + temp[i + "_free_response"]) / 
                      (temp[i + " - Max Points"] + temp[i + "_free_response" + " - Max Points"]))
        else:
            sol[i] = temp[i] / temp[i + " - Max Points"]
        
    return sol.mean(axis=1)

projects_total(grades)
```

Output:
```
/var/folders/s1/y8hjc_rd18g3zw5xh5l4_y8r0000gn/T/ipykernel_80350/1935500803.py:13: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  return sol.mean(axis=1)
0      0.916234
1      0.765932
2      0.681279
3      0.962581
4      0.737446
         ...   
530    0.949434
531    0.866795
532    0.862050
533    0.813468
534    0.939433
Length: 535, dtype: float64
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Computing lab grades

Now, you will clean and process the lab grades, which is a little more complicated. To do this, you will develop functions that:
- 'normalize' the grades, 
- adjust for late submissions, 
- drop the lowest lab grade, and 
- creates a total lab score for each student.

**Question 3**

Unfortunately, Gradescope sometimes experiences a delay in registering when an assignment is submitted during "periods of heavy usage" (i.e. near a submission deadline). You need to assess when a student's assignment was actually turned in on time, even if Gradescope did not process it in time. To do this, it is helpful to know:
* Every late submission has to be submitted by a TA (late submissions are turned off).
* TAs never submitted a late assignment "just after" the deadline. 
* The deadlines were at midnight and students had to come to staff hours to late-submit their assignment.

Create a function `last_minute_submissions` that takes in the dataframe `grades` and outputs the number of submissions on each *lab* assignment that were turned in on time by a student, yet marked 'late' by Gradescope. See the doctest for more details.

*Note:* You have to figure out what truly is a late submission by looking at the data and understanding the facts about the data generating process above. There is some ambiguity in finding which submissions are truly late; you will *make a best guess for a threshold* by looking at this dataset. This question is about 'cleaning' a messy 'data recording process'.

*Note 2:* The return value of your function should only contain counts for the *labs*; other assignment types do not need to be handled.

```python
def last_minute_submissions(grades):
    labs = get_assignment_names(grades)["lab"]    
    sol = {}
    
    for i in labs:
        curr = grades.columns.str.contains(i)
        late = grades.columns[curr]
        
        temp = pd.DataFrame(grades, columns = late)
        temp = temp[temp[i + " - Lateness (H:M:S)"] != "00:00:00"]
        
        # get only under 8 hours
        temp = temp[(temp[i + " - Lateness (H:M:S)"].str.slice(stop = 2)).astype(int) <= 8]
        sol[i] = temp.shape[0]
        
    return pd.Series(sol)

#last_minute_submissions(grades)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4**

Now you need to adjust the lab grades for late submissions -- however, you need to take into account your investigation in the previous question, since students shouldn't be penalized by a bug in Gradescope!

Create a function `lateness_penalty` that takes in a 'Lateness' column and returns a column of penalties (represented by the values `1.0,0.9,0.7,0.4` according to the syllabus). Only *truly* late submissions should be counted as late.

*Note*: For the purpose of this project, we will only be calculating lateness for labs. There is no penalty for lateness for projects, discussions, nor checkpoints.

```python
def lateness_penalty(lateness):
    
    # inner function to apply on entire parameter column
    def inner_lateness(value):
        temp = int(value.split(":")[0])
        if temp == 0:
            return 1.0
        
        if temp < (24 * 7):
            return 0.9
        
        if temp < (24 * 14):
            return 0.7

        return 0.4
    
    
    dupl = lateness
    dupl = dupl.apply(lambda x: "00:00:00" if int(x.split(":")[0]) <= 8 else x)
    return dupl.apply(inner_lateness)

#lateness_penalty(grades["lab06 - Lateness (H:M:S)"])
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Question 5**

Create a function `process_labs` that takes in a dataframe like `grades` and returns a dataframe of processed lab scores. The output should:
* share the same index as `grades`,
* have columns given by the lab assignment names (e.g. `lab01,...lab10`)
* have values representing the lab grades for each assignment, adjusted for Lateness and scaled to a score between 0 and 1.

```python
def process_labs(grades):
    #temp = grades.fillna(0)
    temp = grades
    output = pd.DataFrame()
    for i in get_assignment_names(temp)["lab"]:
        late = lateness_penalty(temp[i + " - Lateness (H:M:S)"])
        output[i] = (temp[i] * late) / temp[i + " - Max Points"]
        
        # clipping scores under 100% and above 0%
        output[i] = np.clip(output[i], 0.00, 1.00)
        
    
    return output

process_labs(grades)
```

Output:
```
lab01     lab02     lab03     lab04     lab05     lab06     lab07  \
0    0.997353  0.849902  0.637744  1.000000  1.000000  0.994518  0.389141   
1    0.988295  0.507842  0.714477  0.783672  1.000000  0.393887  0.914061   
2    0.865134  0.478028  0.433667  0.738875  0.927838  0.345076  0.734070   
3    1.000000  1.000000  0.925903  0.950614  0.891614  0.688403  0.985371   
4    0.665070  0.334224  0.706932  0.747915  0.659720  0.731345  0.607859   
..        ...       ...       ...       ...       ...       ...       ...   
530  0.900000  0.820228  1.000000  0.792935  1.000000  0.284106  0.770281   
531  1.000000  0.874981  0.809945  0.592866  0.987597  0.759688  0.856178   
532  0.886566  0.903260  1.000000  1.000000  0.941425  0.768909  0.967282   
533  0.837997  0.856369  0.909363  0.955287  0.737854  0.382781  0.769093   
534  0.883893  0.730452  0.944795  0.819100  0.890968  0.858277  0.989410   

        lab08     lab09  
0    0.887917  0.874913  
1    0.944378  0.902977  
2    
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Question 6**

Create a function `lab_total` that takes in dataframe of processed assignments (like the output of Question 5) and computes the total lab grade for each student according to the syllabus (returning a Series). Your answers should be proportions between 0 and 1. For example, if there are only 3 labs, and a student received scores of {80%,90%,100%}, then the total score would be 0.95.

*Note*: Don't forget to properly handle students who didn't turn in assignments! (Use your experience and common sense).

```python
def lab_total(processed):
    sol = processed.fillna(0)    
    sol = (sol.sum(axis = 1) - sol.min(axis = 1)) / (len(processed.columns) - 1)
    return sol

lab_total(process_labs(grades))
```

Output:
```
0      0.905293
1      0.844463
2      0.706707
3      0.936836
4      0.686128
         ...   
530    0.901836
531    0.841369
532    0.947054
533    0.860098
534    0.865609
Length: 535, dtype: float64
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

### Putting it all together

**Question 7**

Finally, you need to create the final course grades. To do this, you will add up the total of each course component according to the weights given in the syllabus. 

* Create a function `total_points` that takes in `grades` and returns the final course grades according to the syllabus. Course grades should be proportions between zero and one.
* Create a function `final_grades` that takes in the final course grades as above and returns a Series of letter grades given by the standard cutoffs (`A >= .90`, `.90 > B >= .80`, `.80 > C >= .70`, `.70 > D >= .60`, `.60 > F`). You should not use rounding to determining the letter grades.
* Create a function `letter_proportions` which takes in the dataframe `grades` and outputs a Series that contains the proportion of the class that received each grade. (This question requires you to put everything together).
* The indices should be ordered by the proportion of the class that receives that grade, from largest to smallest.

*Note 1*: Don't repeat yourself when computing the checkpoint and discussion portions of the course.

*Note 2*: Only the lab portion of the course accounts for late assignments; you may assume all assignments in other portions are turned in without penalty.

*Note 3*: These values should add up to exactly 1.0. If you are getting something close such as 0.99999, that means there is a slight issue with your code from above. 

To check your work, verify the course grade distribution and relevant statistics! Do the work by hand for a few students.

```python
def total_points(grades):
    temp = grades.fillna(0)
    sol = pd.DataFrame()
    
    sol["project"] = projects_total(temp)
    sol["lab"] = lab_total(process_labs(temp))
    
    cols = get_assignment_names(temp)
    rem = ["final", "midterm", "disc", "checkpoint"]
    
    for i in rem:
        value = cols[i]
        max_vals = [x + " - Max Points" for x in value]
        sol[i] = temp[value].sum(axis = 1) / temp[max_vals].sum(axis = 1)
        
    final_grade = (sol["lab"] * 0.2) + (sol["project"] * 0.3) + (sol["checkpoint"] * 0.025) +\
    (sol["disc"] * 0.025) + (sol["midterm"] * 0.15) + (sol["final"] * 0.3)
    return final_grade
        
        
#total_points(grades)
```

```python
def final_grades(total):
    def cutoff(val):
        if val >= 0.9:
            return "A"
        elif val >= 0.8:
            return "B"
        elif val >= 0.7:
            return "C"
        elif val >= 0.6:
            return "D"
        else:
            return "F"
    
    # use for checking values
    #for i in range(total.size):
        #print("score: " + str(total[i]) + " and the grade is " + cutoff(total[i]))
        
    sol = total.apply(cutoff)
    return sol

#final_grades(total_points(grades))
```

```python
def letter_proportions(grades):
    temp = total_points(grades)
    temp = final_grades(temp)
    return temp.value_counts(normalize = True)

#letter_proportions(grades)
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

### Do Seniors get worse grades?

**Question 8**

You notice that students who are seniors on average did worse in the class (if you can't verify this, you should go back and check your work!). Is this difference significant, or just due to noise?

Perform a hypothesis test, assessing the likelihood of the above statement under the null hypothesis: 
> "Seniors earn grades that are roughly equal on average to the rest of the class."


Create a function `simulate_pval` which takes in the number of simulations `N` and `grades` and returns the the likelihood that the grade of seniors was worse than the average of the class as a whole under the null hypothesis (i.e. calculate the p-value).

*Note:* To check your work, plot the sampling distribution and the observation. Do these values look reasonable?

*Note 2*: If you sample the data, make sure you sample *with* replacement.

```python
def simulate_pval(grades, N):
    df = pd.concat([grades["Level"], pd.Series(total_points(grades), name = "Total")], axis = 1)
    senior_grades = df[df["Level"] == "SR"].get("Total").mean()
    #check with others
    #change the column name from 0 to smth else
    
    avg = list()
    for i in np.arange(N):
        temp_sample = df["Total"].sample((df[df["Level"] == "SR"]).shape[0], replace = True)
        temp = np.mean(temp_sample)
        avg.append(temp)
        
    return (senior_grades >= pd.Series(avg)).mean()
    

#simulate_pval(grades, 10000)
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

### What is the true distribution of grades?

The gradebook for this class only reflects one particular instance of each student's performance, subject to the effects of all the little events and hiccups that occurred throughout the quarter. Might you have done better on the midterm had your roommate kept you up all night with their coughing? Wasn't it lucky that the example you were studying just before the final happened to appear on the exam?

**Question 9**

This question will simulate these '(un)lucky, random events' by adding or subtracting random amounts to each assignment before calculating the final grades. These 'random amounts' will be drawn from a Gaussian distribution of mean 0 and a std deviation 0.02:
```
np.random.normal(0, 0.02, size=(num_rows, num_cols))
```
Intuitively, such a model says that random events may bump up or down a given grade (given as a proportion):
- which on average has no effect on the class as a whole (mean 0),
- which not uncommonly might perturb a grade by 2% (std dev 0.02).

Create a function `total_points_with_noise` that takes in a dataframe like `grades`, adds noise to the assignments as described above, and returns the final scores using *the same procedure* as questions 1-7.

*Note:* You should be able to reuse (or minorly change) the code from previous problems. Try to be DRY (don't repeat yourself)!

*Note 1:* Once adding the noise to the assignment scores, use the `np.clip` function to be sure each assignment retains a score between 0% and 100%.

*Note 2:* To check your work -- what would you expect the difference between the actual scores and noisy scores to be, on average?

```python
total_points(grades)
```

Output:
```
/var/folders/s1/y8hjc_rd18g3zw5xh5l4_y8r0000gn/T/ipykernel_80350/1935500803.py:13: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  return sol.mean(axis=1)
0      0.902465
1      0.816654
2      0.759665
3      0.908073
4      0.675083
         ...   
530    0.865660
531    0.764832
532    0.859648
533    0.866322
534    0.894719
Length: 535, dtype: float64
```

```python
def total_points_with_noise(grades):
    #using names, we iterate through EACH assignment column in grades
    #for each student's individual assignment, i.e. a cell in the df, 
    #val += np.random.normal(0, 0.02, grades[ass name].shape[0])
    #then at the end we call total_points with this edited df
    
    temp = grades
    names = [j for i in get_assignment_names(grades).values() for j in i]
    
    for i in names:
        noise = np.random.normal(0, 0.02, len(temp[i]))
        temp[i] += (noise * temp[i + " - Max Points"])
        
        #SOMEHOW THIS WORKS
        #temp[i].apply(lambda x: x + np.random.normal(0, 0.02, len(grades[i])))
        
        #INCORRECT ANSWERS BELOW
        #temp[i] += np.random.normal(0, 0.02)
        #temp[i].apply(lambda x: x + np.random.normal(0, 0.02))

        
        #the third parameter makes all the noise values unique, keep it!
        #temp[i] = np.clip(temp[i], 0.00, 100)
        temp[i] = np.clip(temp[i], 0.00, temp[i + " - Max Points"].iloc[0])
                
    return total_points(temp)
    
total_points_with_noise(grades)
```

Output:
```
/var/folders/s1/y8hjc_rd18g3zw5xh5l4_y8r0000gn/T/ipykernel_80350/1935500803.py:13: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  return sol.mean(axis=1)
0      0.898010
1      0.808848
2      0.767833
3      0.919537
4      0.678776
         ...   
530    0.851588
531    0.751379
532    0.863709
533    0.869327
534    0.902596
Length: 535, dtype: float64
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

### Short-answer questions (hard-coded)

Use your functions from above to understanding the data and answer the following questions. The function below should return **hard-coded values**. It should not compute anything!

**Question 10**

Create a function `short_answer` with zero parameters that returns (hard-coded) answers to the following question in a list of length 5:

1. For the class on average, what is the difference between students' scores (`total_points`) and their scores with noise (`total_points_with_noise`)? (Remark: plot the distribution of differences; does this align with what you know about binomial distributions?)
2. What **percentage** of the class only sees their grade change at most (but not including) $\pm 0.01$? (Your answer should be a number between 0 and 100.)
3. What is the 95% confidence interval for the statistic above? (see [DSC10](https://www.inferentialthinking.com/chapters/13/3/Confidence_Intervals.html) and use `np.percentile`)
4. What **proportion** of the class sees a change in their letter grade? (Your answer should be a number between 0 and 1.)
5. Is the assumption behind the model in Question 9 that:
    - The (observed) gradebook well represents the true population of students? (True or False)
    - The noisy scores does not represent other possible observations drawn from the true population of students. (True or False)
    - Answer `True` or `False` in a list, like `[True, True]`.

```python
diff = (abs(total_points(grades) - total_points_with_noise(grades)) > 0.01).sum()
diff
```

Output:
```
/var/folders/s1/y8hjc_rd18g3zw5xh5l4_y8r0000gn/T/ipykernel_80350/1935500803.py:13: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  return sol.mean(axis=1)
87
```

```python
def short_answer():
    # CHECK NUMBER 9, make sure its 100% correct
    
    diff = total_points(grades) - total_points_with_noise(grades)
    q1 = (total_points(grades) - total_points_with_noise(grades)).mean()
    
    q2 = abs(diff)
    q2 = (q2 < 0.01).sum() / len(q2)

    
    vals = np.array([])
    prev = total_points(grades)
    new_gr = total_points_with_noise(grades)
    for i in range(100):
        prev_strap = prev.sample(prev.shape[0], replace = True)
        new_strap = new_gr.sample(new_gr.shape[0], replace = True)
        mic = abs(total_points(grades) - total_points_with_noise(grades))
        per = (np.count_nonzero(mic < 0.01) / mic.shape[0]) * 100
        vals = np.append(vals, per)
        
    q3 = [np.percentile(vals, 2.5), np.percentile(vals, 97.5)]
    
    
    prev = final_grades(total_points(grades))
    new_gr = final_grades(total_points_with_noise(grades))
    q4 = np.sum(prev != new_gr) / prev.shape[0]
    #q4 = q4.apply(ord) - final_grades(total_points(grades)).apply(ord)
    #q4 = q4.where(q4 != 0).shape[0] / q4.shape[0]

    
    #print("Mean is " + str(total_points_with_noise(grades).mean()))
    #print("Min is " + str(total_points_with_noise(grades).min()))
    #print("Max is " + str(total_points_with_noise(grades).max()))
    #print("Median is " + str(total_points_with_noise(grades).median()))
    
    
    q5 = list() 
    q5.append(True) #Because the min/max/median/mean is realistic, and median/mean are around 83%
    q5.append(True) 
    #Thinking about other possible factors outside of noise such as student getting covid 
    #and 2 weeks worth of assignments/labs/projects/etc are low
    
    
    return [q1, q2, q3, q4, q5]

short_answer()
```

Output:
```
/var/folders/s1/y8hjc_rd18g3zw5xh5l4_y8r0000gn/T/ipykernel_80350/1935500803.py:13: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  return sol.mean(axis=1)
[0.0008632870417972022,
 0.8149532710280374,
 [81.39719626168224, 87.57476635514018],
 0.04672897196261682,
 [True, True]]
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

# Congratulations, you finished the project!

### Before you submit:
* Be sure you run the doctests on all your code in `project.py`

### To submit:
* **Upload the .py file to gradescope**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q10 results: All test cases passed!

q2 results:
    q2 - 1 result:
        Test case passed!

    q2 - 2 result:
        Trying:
            out = projects_total(grades)
        Expecting nothing
        ok
        Trying:
            0.7 < out.mean() < 0.9
        Expecting:
            True
        **********************************************************************
        Line 2, in q2 1
        Failed example:
            0.7 < out.mean() < 0.9
        Expected:
            True
        Got:
            False

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results:
    q5 - 1 result:
        Test case passed!

    q5 - 2 result:
        Trying:
            out = process_labs(grades)
        Expecting nothing
        ok
        Trying:
            np.all((0.65 <= out.mean()) & (out.mean() <= 0.90))
        Expecting:
            True
        **********************************************************************
    
```

Tests for q1 (points: None)
```
>>> out = get_assignment_names(grades)
>>> out['final'] == ['Final']
True
```
```
>>> out = get_assignment_names(grades)
>>> set(out.keys()) == {'lab', 'project', 'midterm', 'final', 'disc', 'checkpoint'}
True
```
```
>>> out = get_assignment_names(grades)
>>> 'project02' in out['project']
True
```


Tests for q10 (points: None)
```
>>> out = short_answer()
>>> len(out) == 5
True
```
```
>>> out = short_answer()
>>> (len(out[2])) == 2 and (50 < out[2][0] < 100) and (0 < out[3] < 1)
True
```
```
>>> out = short_answer()
>>> isinstance(out[4][0], bool) and isinstance(out[4][1], bool)
True
```


Tests for q2 (points: None)
```
>>> out = projects_total(grades)
>>> np.all((0 <= out) & (out <= 1))
True
```
```
>>> out = projects_total(grades)
>>> 0.7 < out.mean() < 0.9
True
```


Tests for q3 (points: None)
```
>>> out = last_minute_submissions(grades)
>>> isinstance(out, pd.Series)
True
```
```
>>> out = last_minute_submissions(grades)
>>> np.all(out.index == ['lab0%d' % d for d in range(1,10)])
True
```
failure_message: Checking that all labs, and only labs, are included in the results of last_minute_submissions.
```
>>> out = last_minute_submissions(grades)
>>> (out > 0).sum() == 8
True
```
failure_message: Checking that there are exactly 8 labs that are incorrectly marked as late.


Tests for q4 (points: None)
```
>>> out = lateness_penalty(grades['lab01 - Lateness (H:M:S)'])
>>> isinstance(out, pd.Series)
True
```
```
>>> out = lateness_penalty(grades['lab01 - Lateness (H:M:S)'])
>>> set(out.unique()) <= {1.0, 0.9, 0.7, 0.4}
True
```


Tests for q5 (points: None)
```
>>> out = process_labs(grades)
>>> out.columns.tolist() == ['lab%02d' % x for x in range(1,10)]
True
```
```
>>> out = process_labs(grades)
>>> np.all((0.65 <= out.mean()) & (out.mean() <= 0.90))
True
```


Tests for q6 (points: None)
```
>>> cols = 'lab01 lab02 lab03'.split()
>>> processed = pd.DataFrame([[0.2, 0.90, 1.0]], index=[0], columns=cols)
>>> np.isclose(lab_total(processed), 0.95).all()
True
```


Tests for q7 (points: None)
```
>>> out = total_points(grades)
>>> np.all((0 <= out) & (out <= 1))
True
```
```
>>> out = total_points(grades)
>>> 0.7 < out.mean() < 0.9
True
```
```
>>> out = final_grades(pd.Series([0.92, 0.81, 0.41]))
>>> np.all(out == ['A', 'B', 'F'])
True
```
```
>>> out = letter_proportions(grades)
>>> np.all(out.index == ['B', 'C', 'A', 'D', 'F'])
True
```
```
>>> out = letter_proportions(grades)
>>> out.sum() == 1.0
True
```


Tests for q8 (points: None)
```
>>> out = simulate_pval(grades, 1000)
>>> 0 <= out <= 0.1
True
```


Tests for q9 (points: None)
```
>>> out = total_points_with_noise(grades)
>>> np.all((0 <= out) & (out <= 1))
True
```
```
>>> out = total_points_with_noise(grades)
>>> 0.7 < out.mean() < 0.9
True
```
